<!DOCTYPE html>


<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    
    <title>7.9. Exercises &mdash; Runestone Interactive Overview</title>
    
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="../_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="../_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="../_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="../_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../_static/livecode.css" type="text/css" />
    <link rel="stylesheet" href="../_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="../_static/lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="../_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="../_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="../_static/video.css" type="text/css" />
    <link rel="stylesheet" href="http://netdna.bootstrapcdn.com/bootswatch/2.3.1/""/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    
    <script type="text/javascript">
      var DOCUMENTATION_OPTIONS = {
        URL_ROOT:    '../',
        VERSION:     '0.0',
        COLLAPSE_INDEX: false,
        FILE_SUFFIX: '.html',
        HAS_SOURCE:  true
      };
    </script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/livecode.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="https://cdn.mathjax.org/mathjax/latest/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="top" title="Runestone Interactive Overview" href="../index.html" />
    <link rel="up" title="7. Functional List Processing" href="toctree.html" />
    <link rel="next" title="Lab 1 - Closures and Scope" href="../Labs/ClosuresAndScope.html" />
    <link rel="prev" title="7.8. Working with Large Files" href="WorkingWithLargeFiles.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'STAT489';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'STAT489';
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>



  </head>
  <body role="document">


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html">STAT489</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-share-alt" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/STAT489/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/STAT489/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=STAT489&page=Exercises'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">7.9. Exercises</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - 7.8. Working with Large Files" data-toggle="tooltip">
    <a href="WorkingWithLargeFiles.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Lab 1 - Closures and Scope' data-toggle="tooltip" >
    <a href="../Labs/ClosuresAndScope.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="exercises">
<h1>7.9. Exercises<a class="headerlink" href="#exercises" title="Permalink to this headline">¶</a></h1>

    <div data-component="question" class="full-width container question" id="15" >
    <ol start=1 class=arabic><li class="alert alert-warning">

    <p>In Robert McCloskey&#8217;s book <em>Make Way for Ducklings</em>, the names of the
ducklings are Jack, Kack, Lack, Mack, Nack, Ouack, Pack, and Quack.  This
loop tries to output these names in order. Save your answer in a variable
named <code class="docutils literal"><span class="pre">words</span></code> and make sure that this value doesn&#8217;t get overwritten later in
your program.</p>
<p>Use map plus a helper function to solve this problem.</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">prefixes</span> <span class="o">=</span> <span class="s2">&quot;JKLMNOPQ&quot;</span>
<span class="n">suffix</span> <span class="o">=</span> <span class="s2">&quot;ack&quot;</span>

<span class="n">words</span> <span class="o">=</span> <span class="p">[</span><span class="n">first</span> <span class="o">+</span> <span class="n">suffix</span> <span class="k">for</span> <span class="n">first</span> <span class="ow">in</span> <span class="n">prefixes</span><span class="p">]</span>
</pre></div>
</div>
<p>Of course, that&#8217;s not quite right because Ouack and Quack are misspelled.
Can you fix it?</p>
<blockquote>
<div>
<textarea data-component="activecode" id=ex_8_2 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >


</textarea>
</div></blockquote>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="486" >
    <ol start=2 class=arabic><li class="alert alert-warning">

    <p>Use map and filter to filter the hours table to include only managers.
In SQL this would be performed using SELECT and WHERE.  <strong>Hint:</strong> Start by
creating a list of the names of all managers.</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">hours</span> <span class="o">=</span> <span class="p">[[</span><span class="s2">&quot;Alice&quot;</span><span class="p">,</span> <span class="mi">43</span><span class="p">],</span>
           <span class="p">[</span><span class="s2">&quot;Bob&quot;</span><span class="p">,</span> <span class="mi">37</span><span class="p">],</span>
           <span class="p">[</span><span class="s2">&quot;Fred&quot;</span><span class="p">,</span> <span class="mi">15</span><span class="p">]]</span>
<span class="n">titles</span> <span class="o">=</span> <span class="p">[[</span><span class="s2">&quot;Alice&quot;</span><span class="p">,</span> <span class="s2">&quot;Manager&quot;</span><span class="p">],</span>
          <span class="p">[</span><span class="s2">&quot;Betty&quot;</span><span class="p">,</span> <span class="s2">&quot;Consultant&quot;</span><span class="p">],</span>
          <span class="p">[</span><span class="s2">&quot;Bob&quot;</span><span class="p">,</span> <span class="s2">&quot;Assistant&quot;</span><span class="p">]]</span>
</pre></div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="502" >
    <ol start=3 class=arabic><li class="alert alert-warning">

    <p>Use map, len, and filter to decide if the following tables contain a manager
that worked at least 40 hours.</p>
<div class="highlight-python"><div class="highlight"><pre><span></span><span class="n">hours</span> <span class="o">=</span> <span class="p">[[</span><span class="s2">&quot;Alice&quot;</span><span class="p">,</span> <span class="mi">43</span><span class="p">],</span>
           <span class="p">[</span><span class="s2">&quot;Bob&quot;</span><span class="p">,</span> <span class="mi">37</span><span class="p">],</span>
           <span class="p">[</span><span class="s2">&quot;Fred&quot;</span><span class="p">,</span> <span class="mi">15</span><span class="p">]]</span>
<span class="n">titles</span> <span class="o">=</span> <span class="p">[[</span><span class="s2">&quot;Alice&quot;</span><span class="p">,</span> <span class="s2">&quot;Manager&quot;</span><span class="p">],</span>
          <span class="p">[</span><span class="s2">&quot;Betty&quot;</span><span class="p">,</span> <span class="s2">&quot;Consultant&quot;</span><span class="p">],</span>
          <span class="p">[</span><span class="s2">&quot;Bob&quot;</span><span class="p">,</span> <span class="s2">&quot;Assistant&quot;</span><span class="p">]]</span>
</pre></div>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="531" >
    <ol start=4 class=arabic><li class="alert alert-warning">

    <p>Use some combination of map and filter to create a sequence of
functions that combine to average two matrices.  A complete solution will
provide functions for each level of abstraction.</p>

<textarea data-component="activecode" id=average-matrices data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
M1 = [[1, 2], [3, 4]]
M2 = [[5, 6], [7, 8]]
</textarea>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="452" >
    <ol start=5 class=arabic><li class="alert alert-warning">

    <div id='q3' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Create a function that takes a value <cite>n</cite> as input and constructs a
multiplication table for whole numbers up to <span class="math">\(n\)</span>.
Use some combination of map, filter and reduce to solve this problem.</p>

<textarea data-component="activecode" id=ex_8_3 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >


</textarea>
</div><div data-component="tab" data-tabname="Answer" >
<div class="highlight-ipython"><div class="highlight"><pre><span></span><span class="gp">In [1]: </span><span class="n">row</span> <span class="o">=</span> <span class="k">lambda</span> <span class="n">j</span><span class="p">,</span> <span class="n">n</span><span class="p">:</span> <span class="nb">list</span><span class="p">(</span><span class="nb">map</span><span class="p">(</span><span class="k">lambda</span> <span class="n">i</span><span class="p">:</span> <span class="n">j</span><span class="o">*</span><span class="n">i</span><span class="p">,</span> <span class="nb">range</span><span class="p">(</span><span class="mi">1</span><span class="p">,</span> <span class="n">n</span><span class="o">+</span><span class="mi">1</span><span class="p">)))</span>

<span class="gp">In [2]: </span><span class="n">table</span> <span class="o">=</span> <span class="k">lambda</span> <span class="n">n</span><span class="p">:</span> <span class="nb">list</span><span class="p">(</span><span class="nb">map</span><span class="p">(</span><span class="k">lambda</span> <span class="n">j</span><span class="p">:</span> <span class="n">row</span><span class="p">(</span><span class="n">j</span><span class="p">,</span> <span class="n">n</span><span class="p">),</span> <span class="nb">range</span><span class="p">(</span><span class="mi">1</span><span class="p">,</span> <span class="n">n</span><span class="o">+</span><span class="mi">1</span><span class="p">)))</span>

<span class="gp">In [3]: </span><span class="n">table</span><span class="p">(</span><span class="mi">12</span><span class="p">)</span>
<span class="gh">Out[3]: </span><span class="go"></span>
<span class="go">[[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12],</span>
<span class="go"> [2, 4, 6, 8, 10, 12, 14, 16, 18, 20, 22, 24],</span>
<span class="go"> [3, 6, 9, 12, 15, 18, 21, 24, 27, 30, 33, 36],</span>
<span class="go"> [4, 8, 12, 16, 20, 24, 28, 32, 36, 40, 44, 48],</span>
<span class="go"> [5, 10, 15, 20, 25, 30, 35, 40, 45, 50, 55, 60],</span>
<span class="go"> [6, 12, 18, 24, 30, 36, 42, 48, 54, 60, 66, 72],</span>
<span class="go"> [7, 14, 21, 28, 35, 42, 49, 56, 63, 70, 77, 84],</span>
<span class="go"> [8, 16, 24, 32, 40, 48, 56, 64, 72, 80, 88, 96],</span>
<span class="go"> [9, 18, 27, 36, 45, 54, 63, 72, 81, 90, 99, 108],</span>
<span class="go"> [10, 20, 30, 40, 50, 60, 70, 80, 90, 100, 110, 120],</span>
<span class="go"> [11, 22, 33, 44, 55, 66, 77, 88, 99, 110, 121, 132],</span>
<span class="go"> [12, 24, 36, 48, 60, 72, 84, 96, 108, 120, 132, 144]]</span>
</pre></div>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="474" >
    <ol start=6 class=arabic><li class="alert alert-warning">

    Print out a neatly formatted multiplication table, up to 12 x 12.  You should
do this by constructing a string.  For full credit, each column should be
right-justified and your solution should include only <code class="docutils literal"><span class="pre">map</span></code> and helper
functions. (You are allowed to use the strings <code class="docutils literal"><span class="pre">join</span></code> and <code class="docutils literal"><span class="pre">rjust</span></code> methods
in your helper functions) <strong>Hint:</strong> Write a lambda function pads a number with the right
number of spaces.  Use some combination of map, filter and reduce to solve
this problem.
    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="37" >
    <ol start=7 class=arabic><li class="alert alert-warning">

    <div id='q5' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Write a function that will return the number of digits in an integer.</p>
<p>Use reduce to solve thie problem.</p>

<textarea data-component="activecode" id=ex_7_10 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
numDigits = lambda n: 0  # your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(numDigits(2),1,"Tested numDigits on input of 2")
      self.assertEqual(numDigits(55),2,"Tested numDigits on input of 55")
      self.assertEqual(numDigits(1352),4,"Tested numDigits on input of 1352")
      self.assertEqual(numDigits(444),3,"Tested numDigits on input of 444")



myTests().main()
</textarea>
</div><div data-component="tab" data-tabname="Answer" >
<div class="highlight-ipython"><div class="highlight"><pre><span></span><span class="gp">In [4]: </span><span class="kn">from</span> <span class="nn">functools</span> <span class="kn">import</span> <span class="nb">reduce</span>

<span class="gp">In [5]: </span><span class="n">numDigits</span> <span class="o">=</span> <span class="k">lambda</span> <span class="n">n</span><span class="p">:</span> <span class="nb">reduce</span><span class="p">(</span><span class="k">lambda</span> <span class="n">acc</span><span class="p">,</span> <span class="n">ch</span><span class="p">:</span> <span class="n">acc</span> <span class="o">+</span> <span class="mi">1</span><span class="p">,</span> <span class="nb">str</span><span class="p">(</span><span class="n">n</span><span class="p">),</span> <span class="mi">0</span><span class="p">)</span>

<span class="gp">In [6]: </span><span class="n">numDigits</span><span class="p">(</span><span class="mi">50</span><span class="p">)</span>
<span class="gh">Out[6]: </span><span class="go">2</span>

<span class="gp">In [7]: </span><span class="n">numDigits</span><span class="p">(</span><span class="mi">20000</span><span class="p">)</span>
<span class="go">          Out[7]: 5</span>

<span class="gp">In [8]: </span><span class="n">numDigits</span><span class="p">(</span><span class="mi">1</span><span class="p">)</span>
<span class="go">                    Out[8]: 1</span>
</pre></div>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="148" >
    <ol start=8 class=arabic><li class="alert alert-warning">

    <p>Write a function named <code class="docutils literal"><span class="pre">remove_char</span></code> that removes all occurrences of a
given letter from a string.  Do not use the <code class="docutils literal"><span class="pre">replace</span></code> string method.</p>
<p>Use reduce to solve this problem.</p>

<textarea data-component="activecode" id=ex_8_7 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >
remove_char = lambda theLetter, theString: "stuff" # your code here

====


from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(remove_char("a","apple"),"pple","Tested remove_char on inputs of 'a' and 'apple'")
      self.assertEqual(remove_char("a","banana"),"bnn","Tested remove_char on inputs of 'a' and 'banana'")
      self.assertEqual(remove_char("z","banana"),"banana","Tested remove_char on inputs of 'z' and 'banana'")



myTests().main()
</textarea>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="119" >
    <ol start=9 class=arabic><li class="alert alert-warning">

    Write a function called <code class="docutils literal"><span class="pre">str_join</span></code> that takes two arguments <code class="docutils literal"><span class="pre">s</span></code> (a
string) and <code class="docutils literal"><span class="pre">L</span></code> (a list of strings) and returns the equivalent of
<code class="docutils literal"><span class="pre">s.join(L)</span></code>.  Use reduce to solve this problem without using the string
<code class="docutils literal"><span class="pre">join</span></code> method.
    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="176" >
    <ol start=10 class=arabic><li class="alert alert-warning">

    <div id='176' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Write a function that recognizes palindromes. Use reduce to solve
thie problem.  <strong>Hint:</strong> Use reversed and zip.</p>

<textarea data-component="activecode" id=ex_8_8 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >
is_palindrome lambda myStr: False # your code here

====


from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(is_palindrome("abba"),True,"Tested is_palindrome on input of 'abba'")
        self.assertEqual(is_palindrome("abab"),False,"Tested is_palindrome on input of 'abab'")
        self.assertEqual(is_palindrome("straw warts"),True,"Tested is_palindrome on input of 'straw warts'")
        self.assertEqual(is_palindrome("a"),True,"Tested is_palindrome on input of 'a'")
        self.assertEqual(is_palindrome(""),True,"Tested is_palindrome on input of ''")




myTests().main()
</textarea>
</div><div data-component="tab" data-tabname="Answer" >
<div class="highlight-ipython"><div class="highlight"><pre><span></span><span class="gp">In [9]: </span><span class="kn">from</span> <span class="nn">functools</span> <span class="kn">import</span> <span class="nb">reduce</span>

<span class="gp">In [10]: </span><span class="n">is_palindrome</span> <span class="o">=</span> <span class="k">lambda</span> <span class="n">mystr</span><span class="p">:</span> <span class="nb">reduce</span><span class="p">(</span><span class="k">lambda</span> <span class="n">acc</span><span class="p">,</span> <span class="n">tup</span><span class="p">:</span> <span class="n">acc</span> <span class="ow">and</span> <span class="n">tup</span><span class="p">[</span><span class="mi">0</span><span class="p">]</span> <span class="o">==</span> <span class="n">tup</span><span class="p">[</span><span class="mi">1</span><span class="p">],</span> <span class="nb">zip</span><span class="p">(</span><span class="n">mystr</span><span class="p">,</span><span class="nb">reversed</span><span class="p">(</span><span class="n">mystr</span><span class="p">)),</span> <span class="bp">True</span><span class="p">)</span>

<span class="gp">In [11]: </span><span class="n">is_palindrome</span><span class="p">(</span><span class="s1">&#39;abba&#39;</span><span class="p">)</span>
<span class="gh">Out[11]: </span><span class="go">True</span>
</pre></div>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="170" >
    <ol start=11 class=arabic><li class="alert alert-warning">

    Write a function called <code class="docutils literal"><span class="pre">has_vowel</span></code> that takes a string as input and returns
<code class="docutils literal"><span class="pre">True</span></code> if the string contains a vowel (<code class="docutils literal"><span class="pre">'aeiou'</span></code> ignore <code class="docutils literal"><span class="pre">'y'</span></code>) and <code class="docutils literal"><span class="pre">False</span></code>
otherwise.  Use <code class="docutils literal"><span class="pre">reduce</span></code> to perform this task.
    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="275" >
    <ol start=12 class=arabic><li class="alert alert-warning">

    <div id='275' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Write a Python function that takes <code class="docutils literal"><span class="pre">n</span></code> and <code class="docutils literal"><span class="pre">m</span></code> as input and
returns the maximum value of a the list of <code class="docutils literal"><span class="pre">n</span></code> random integers
between 0 and <code class="docutils literal"><span class="pre">m</span></code>.  (Note:.  there is a builtin function named
<code class="docutils literal"><span class="pre">max</span></code>.)</p>
<p>Use map to generating the random sequence and reduce to find the max.</p>

<textarea data-component="activecode" id=ex_9_5 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >


</textarea>
</div><div data-component="tab" data-tabname="Answer" >
<div class="highlight-ipython"><div class="highlight"><pre><span></span><span class="gp">In [12]: </span><span class="kn">from</span> <span class="nn">random</span> <span class="kn">import</span> <span class="n">randint</span>

<span class="gp">In [13]: </span><span class="kn">from</span> <span class="nn">functools</span> <span class="kn">import</span> <span class="nb">reduce</span>

<span class="gp">In [14]: </span><span class="kn">from</span> <span class="nn">itertools</span> <span class="kn">import</span> <span class="n">repeat</span>

<span class="gp">In [15]: </span><span class="n">rand_seq</span> <span class="o">=</span> <span class="k">lambda</span> <span class="n">n</span><span class="p">,</span> <span class="n">m</span><span class="p">:</span> <span class="nb">map</span><span class="p">(</span><span class="n">randint</span><span class="p">,</span> <span class="n">repeat</span><span class="p">(</span><span class="mi">0</span><span class="p">,</span> <span class="n">n</span><span class="p">),</span> <span class="n">repeat</span><span class="p">(</span><span class="n">m</span><span class="p">,</span> <span class="n">n</span><span class="p">))</span>

<span class="gp">In [16]: </span><span class="n">update_max</span> <span class="o">=</span> <span class="k">lambda</span> <span class="n">cur_max</span><span class="p">,</span> <span class="n">item</span><span class="p">:</span> <span class="n">item</span> <span class="k">if</span> <span class="n">item</span> <span class="o">&gt;</span> <span class="n">cur_max</span> <span class="k">else</span> <span class="n">cur_max</span>

<span class="gp">In [17]: </span><span class="n">rand_max</span> <span class="o">=</span> <span class="k">lambda</span> <span class="n">n</span><span class="p">,</span> <span class="n">m</span><span class="p">:</span> <span class="nb">reduce</span><span class="p">(</span><span class="n">update_max</span><span class="p">,</span> <span class="n">rand_seq</span><span class="p">(</span><span class="n">n</span><span class="p">,</span> <span class="n">m</span><span class="p">))</span>

<span class="gp">In [18]: </span><span class="n">rand_max</span><span class="p">(</span><span class="mi">10</span><span class="p">,</span><span class="mi">5</span><span class="p">)</span>
<span class="gh">Out[18]: </span><span class="go">4</span>
</pre></div>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="299" >
    <ol start=13 class=arabic><li class="alert alert-warning">

    <p>Write a function called <code class="docutils literal"><span class="pre">mean_normal(n,</span> <span class="pre">m,</span> <span class="pre">sd)</span></code> that computes the mean of <code class="docutils literal"><span class="pre">n</span></code> randomly selected values
taken from a normal distribution with mean <code class="docutils literal"><span class="pre">m</span></code> and standard deviation
<code class="docutils literal"><span class="pre">sd</span></code>.  Use map and normalvariate (from the random module) to generate the
sequence and use reduce to compute the mean.</p>

<textarea data-component="activecode" id=ex_9_4 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >


</textarea>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="310" >
    <ol start=14 class=arabic><li class="alert alert-warning">

    <p>Write a function <code class="docutils literal"><span class="pre">sum_of_squares(xs)</span></code> that computes the sum
of the squares of the numbers in the list <code class="docutils literal"><span class="pre">xs</span></code>.  For example,
<code class="docutils literal"><span class="pre">sum_of_squares([2,</span> <span class="pre">3,</span> <span class="pre">4])</span></code> should return 4+9+16 which is 29.</p>
<p>Use reduce to solve this problem.</p>

<textarea data-component="activecode" id=ex_7_11 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
sum_of_squares = lambda xs: 1 # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sum_of_squares([2,3,4]),29,"Tested sum_of_squares on input [2,3,4]")
        self.assertEqual(sum_of_squares([0,1,-1]),2,"Tested sum_of_squares on input [0,1,-1]")
        self.assertEqual(sum_of_squares([5,12,14]),365,"Tested sum_of_squares on input [5,12,14]")

myTests().main()
</textarea>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="334" >
    <ol start=15 class=arabic><li class="alert alert-warning">

    <div id='q7' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Write a function to count how many odd numbers are in a list.</p>
<p>Use some combination of map, filter and reduce to solve this problem.</p>

<textarea data-component="activecode" id=ex_9_6 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
countOdd = lambda lst: 3 # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(countOdd([1,3,5,7,9]),5,"Tested countOdd on input [1,3,5,7,9]")
        self.assertEqual(countOdd([1,2,3,4,5]),3,"Tested countOdd on input [-1,-2,-3,-4,-5]")
        self.assertEqual(countOdd([2,4,6,8,10]),0,"Tested countOdd on input [2,4,6,8,10]")
        self.assertEqual(countOdd([0,-1,12,-33]),2,"Tested countOdd on input [0,-1,12,-33]")

myTests().main()
</textarea>
</div><div data-component="tab" data-tabname="Answer" >
<div class="highlight-ipython"><div class="highlight"><pre><span></span><span class="gp">In [19]: </span><span class="kn">import</span> <span class="nn">random</span>

<span class="gp">In [20]: </span><span class="n">countOdd</span> <span class="o">=</span> <span class="k">lambda</span> <span class="n">lst</span><span class="p">:</span> <span class="nb">reduce</span><span class="p">(</span><span class="k">lambda</span> <span class="n">a</span><span class="p">,</span> <span class="n">i</span><span class="p">:</span> <span class="n">a</span> <span class="o">+</span> <span class="mi">1</span><span class="p">,</span> <span class="nb">filter</span><span class="p">(</span><span class="k">lambda</span> <span class="n">n</span><span class="p">:</span> <span class="n">n</span> <span class="o">%</span> <span class="mi">2</span> <span class="o">==</span> <span class="mi">1</span><span class="p">,</span> <span class="n">lst</span><span class="p">))</span>

<span class="go"># make a random list to test the function</span>
<span class="gp">In [21]: </span><span class="n">lst</span> <span class="o">=</span> <span class="p">[</span><span class="n">random</span><span class="o">.</span><span class="n">randint</span><span class="p">(</span><span class="mi">0</span><span class="p">,</span> <span class="mi">1000</span><span class="p">)</span> <span class="k">for</span> <span class="n">i</span> <span class="ow">in</span> <span class="nb">range</span><span class="p">(</span><span class="mi">100</span><span class="p">)]</span>

<span class="gp">In [22]: </span><span class="n">countOdd</span><span class="p">(</span><span class="n">lst</span><span class="p">)</span>
<span class="gh">Out[22]: </span><span class="go">194</span>
</pre></div>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="377" >
    <ol start=16 class=arabic><li class="alert alert-warning">

    <p>Write a function called <code class="docutils literal"><span class="pre">sum_even</span></code> that sums up all the even numbers in a list.
Use some combination of map, filter and reduce to solve this problem.</p>

<textarea data-component="activecode" id=ex_9_7 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
sum_even = lambda lst: 42 # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sum_even([1,3,5,7,9]),0,"Tested sum_even on input [1,3,5,7,9]")
        self.assertEqual(sum_even([-1,-2,-3,-4,-5]),-6,"Tested sum_even on input [-1,-2,-3,-4,-5]")
        self.assertEqual(sum_even([2,4,6,7,9]),12,"Tested sum_even on input [2,4,6,7,9]")
        self.assertEqual(sum_even([0,1,12,33]),12,"Tested sum_even on input [0,1,12,33]")

myTests().main()
</textarea>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="400" >
    <ol start=17 class=arabic><li class="alert alert-warning">

    <div id='q9' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Sum up all the negative numbers in a list.
Use some combination of map, filter and reduce to solve this problem.</p>

<textarea data-component="activecode" id=ex_9_8 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
sumNegatives = lambda lst: -1 # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sumNegatives([-1,-2,-3,-4,-5]),-15,"Tested sumNegatives on input [-1,-2,-3,-4,-5]")
        self.assertEqual(sumNegatives([1,-3,5,-7,9]),-10,"Tested sumNegatives on input [1,-3,5,-7,9]")
        self.assertEqual(sumNegatives([-2,-4,6,-7,9]),-13,"Tested sumNegatives on input [-2,-4,6,-7,9]")
        self.assertEqual(sumNegatives([0,1,2,3,4]),0,"Tested sumNegatives on input [0,1,2,3,4]")

myTests().main()
</textarea>
</div><div data-component="tab" data-tabname="Answer" >
<div class="highlight-ipython"><div class="highlight"><pre><span></span><span class="gp">In [23]: </span><span class="kn">import</span> <span class="nn">random</span>

<span class="gp">In [24]: </span><span class="kn">from</span> <span class="nn">operator</span> <span class="kn">import</span> <span class="n">add</span>

<span class="gp">In [25]: </span><span class="n">sumNegatives</span> <span class="o">=</span> <span class="k">lambda</span> <span class="n">lst</span><span class="p">:</span> <span class="nb">reduce</span><span class="p">(</span><span class="n">add</span><span class="p">,</span> <span class="nb">filter</span><span class="p">(</span><span class="k">lambda</span> <span class="n">n</span><span class="p">:</span> <span class="n">n</span> <span class="o">&lt;</span> <span class="mi">0</span><span class="p">,</span> <span class="n">lst</span><span class="p">))</span>

<span class="gp">In [26]: </span><span class="n">lst</span> <span class="o">=</span> <span class="p">[</span><span class="n">random</span><span class="o">.</span><span class="n">randint</span><span class="p">(</span><span class="o">-</span><span class="mi">1000</span><span class="p">,</span> <span class="mi">1000</span><span class="p">)</span> <span class="k">for</span> <span class="n">i</span> <span class="ow">in</span> <span class="nb">range</span><span class="p">(</span><span class="mi">100</span><span class="p">)]</span>

<span class="gp">In [27]: </span><span class="n">sumNegatives</span><span class="p">(</span><span class="n">lst</span><span class="p">)</span>
<span class="gh">Out[27]: </span><span class="go">-24819</span>
</pre></div>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="442" >
    <ol start=18 class=arabic><li class="alert alert-warning">

    <p>Write a function called <code class="docutils literal"><span class="pre">num_greater_5</span></code> that counts how many words in a list have a more than than 5 characters.
Use some combination of map, filter and reduce to solve this problem.</p>

<textarea data-component="activecode" id=ex_9_9 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
num_greater_5 = lambda lst: 23 # your code here
</textarea>

<textarea data-component="activecode" id=ex_8_4 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >


</textarea>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="files_ex_3_2" >
    <ol start=19 class=arabic><li class="alert alert-warning">

    <div id='q3' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>The following sample file called <code class="docutils literal"><span class="pre">studentdata.txt</span></code> contains one
line for each student in an imaginary class.  The student&#8217;s name is
the first thing on each line, followed by some exam scores.  The
number of scores might be different for each student.</p>

<pre data-component="datafile" id=studentdata.txt  data-edit="false" data-rows="20" data-cols="36">
joe 10 15 20 30 40
bill 23 16 19 22
sue 8 22 17 14 32 17 24 21 2 9 11 17
grace 12 28 21 45 26 10
john 14 32 25 16 89
</pre>
<p>Using the text file <code class="docutils literal"><span class="pre">studentdata.txt</span></code> (shown in above) write
a program that creates a new data table that contains the students&#8217;
names along with the minimum and maximum score for each student.</p>
<p>Complete this exercise using <code class="docutils literal"><span class="pre">map</span></code>, <code class="docutils literal"><span class="pre">filter</span></code> and/or reduce.
<strong>Note:</strong> For full credit you will need to also implement <code class="docutils literal"><span class="pre">max</span></code>
and <code class="docutils literal"><span class="pre">min</span></code> using <code class="docutils literal"><span class="pre">reduce</span></code>.</p>

<textarea data-component="activecode" id=ex_6_3 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio='studentdata.txt*atype*'    data-gradebutton=true >


</textarea>
</div><div data-component="tab" data-tabname="Answer" >

<textarea data-component="activecode" id=ch_files_q3answer data-lang="python"    data-timelimit=25000   data-audio=''     >
# Convert the following solution to one that uses map, filter and/or reduce
# For this exercise,
with open("studentdata.txt", "r") as f:
    table = [line.split() for line in f]
process_row = lambda row: (row[0], max(row), min(row))
new_table = [process_row(row) for row in table]
new_table
</textarea>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="files_ex_4_2" >
    <ol start=20 class=arabic><li class="alert alert-warning">

    <p><a class="reference external" href="http://seanlahman.com/files/database/baseballdatabank-master_2016-03-02.zip">SeanLahman.com</a>
provides a database of baseball statistics.  Download, unzip and extract the
file titled <strong>Batting.csv</strong>.</p>
<p>Write each of the following functions and apply them to the related task.
Write all of your functions and expressions only using funcitonal tools like
<code class="docutils literal"><span class="pre">with_iter</span></code>, <code class="docutils literal"><span class="pre">map</span></code>, <code class="docutils literal"><span class="pre">filter</span></code>, <code class="docutils literal"><span class="pre">reduce</span></code>, <code class="docutils literal"><span class="pre">compose</span></code> and/or <code class="docutils literal"><span class="pre">pipe</span></code>.</p>
<ol class="arabic simple">
<li>Write a function that takes a year and table of batting data as input
and returns the average number of runs scored in that year.  Call
this function <strong>average_runs_year</strong>.</li>
<li>Write a function takes a list of years and a table of batting data and
returns of list of tuples of the form <cite>(year, averages_runs)</cite>.  Call
this function <strong>average_runs_years</strong> Use a list comprehension
and the function from the last step.</li>
<li>Write a program that includes the above functions and reads
<strong>Batting.csv</strong> and compute the average number of runs for the
following years: 1900, 1910, ..., 2000, 2010.  Assume that the
Batting.csv file is the same directory as your program.</li>
</ol>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="complement" >
    <ol start=21 class=arabic><li class="alert alert-warning">

    <p>The <code class="docutils literal"><span class="pre">toolz</span></code> library includes a function called <code class="docutils literal"><span class="pre">complement</span></code>, which takes a
Boolean function as input and returns a Boolean function that is the logical
opposite.</p>
<p>For example, <code class="docutils literal"><span class="pre">complement(is_odd)</span></code> would be equivalent to
<code class="docutils literal"><span class="pre">is_even</span></code> and <code class="docutils literal"><span class="pre">complement(less_than_5)</span></code> would be equivalent to
<code class="docutils literal"><span class="pre">at_least_5</span></code>.</p>
<div class="highlight-ipython"><div class="highlight"><pre><span></span><span class="gp">In [28]: </span><span class="kn">from</span> <span class="nn">toolz</span> <span class="kn">import</span> <span class="n">complement</span>

<span class="gp">In [29]: </span><span class="n">is_odd</span> <span class="o">=</span> <span class="k">lambda</span> <span class="n">n</span><span class="p">:</span> <span class="n">n</span> <span class="o">%</span> <span class="mi">2</span> <span class="o">==</span> <span class="mi">1</span>

<span class="gp">In [30]: </span><span class="n">is_odd</span><span class="p">(</span><span class="mi">5</span><span class="p">)</span>
<span class="gh">Out[30]: </span><span class="go">True</span>

<span class="gp">In [31]: </span><span class="n">is_odd</span><span class="p">(</span><span class="mi">4</span><span class="p">)</span>
<span class="go">              Out[31]: False</span>

<span class="gp">In [32]: </span><span class="n">is_even</span> <span class="o">=</span> <span class="n">complement</span><span class="p">(</span><span class="n">is_odd</span><span class="p">)</span>

<span class="gp">In [33]: </span><span class="n">is_even</span><span class="p">(</span><span class="mi">5</span><span class="p">)</span>
<span class="gh">Out[33]: </span><span class="go">False</span>

<span class="gp">In [34]: </span><span class="n">is_even</span><span class="p">(</span><span class="mi">4</span><span class="p">)</span>
<span class="go">               Out[34]: True</span>
</pre></div>
</div>
<div class="highlight-ipython"><div class="highlight"><pre><span></span><span class="gp">In [35]: </span><span class="n">less_than_5</span> <span class="o">=</span> <span class="k">lambda</span> <span class="n">n</span><span class="p">:</span> <span class="n">n</span> <span class="o">&lt;</span> <span class="mi">5</span>

<span class="gp">In [36]: </span><span class="n">less_than_5</span><span class="p">(</span><span class="mi">2</span><span class="p">)</span>
<span class="gh">Out[36]: </span><span class="go">True</span>

<span class="gp">In [37]: </span><span class="n">less_than_5</span><span class="p">(</span><span class="mi">22</span><span class="p">)</span>
<span class="go">              Out[37]: False</span>

<span class="gp">In [38]: </span><span class="n">at_least_5</span> <span class="o">=</span> <span class="n">complement</span><span class="p">(</span><span class="n">less_than_5</span><span class="p">)</span>

<span class="gp">In [39]: </span><span class="n">at_least_5</span><span class="p">(</span><span class="mi">2</span><span class="p">)</span>
<span class="gh">Out[39]: </span><span class="go">False</span>

<span class="gp">In [40]: </span><span class="n">at_least_5</span><span class="p">(</span><span class="mi">22</span><span class="p">)</span>
<span class="go">               Out[40]: True</span>
</pre></div>
</div>
<p>Implement your own version of this function named
<code class="docutils literal"><span class="pre">my_complement</span></code>.</p>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="do" >
    <ol start=22 class=arabic><li class="alert alert-warning">

    <p>The <code class="docutils literal"><span class="pre">toolz</span></code> library includes a function called <code class="docutils literal"><span class="pre">do</span></code> that is useful when
performing side effects.  We saw <code class="docutils literal"><span class="pre">do</span></code> in action in Chapter 5 when printing a
sequence out to a file.  The function <code class="docutils literal"><span class="pre">do(func,</span> <span class="pre">x)</span></code> takes a (likely
side-effecting) function <code class="docutils literal"><span class="pre">func</span></code> and a value <code class="docutils literal"><span class="pre">x</span></code> as inputs.  First, <code class="docutils literal"><span class="pre">func</span></code>
is called with <code class="docutils literal"><span class="pre">x</span></code> as an argument then the original <code class="docutils literal"><span class="pre">x</span></code> is returned.  Note
that the any value returned by <code class="docutils literal"><span class="pre">func(x)</span></code> is discarded.</p>
<p>For example, we can use <code class="docutils literal"><span class="pre">do</span></code> to say hello to a given value (with print)
and then return that value.</p>
<div class="highlight-ipython"><div class="highlight"><pre><span></span><span class="gp">In [41]: </span><span class="kn">from</span> <span class="nn">toolz</span> <span class="kn">import</span> <span class="n">do</span>

<span class="gp">In [42]: </span><span class="n">say_hello</span> <span class="o">=</span> <span class="k">lambda</span> <span class="n">x</span><span class="p">:</span> <span class="k">print</span><span class="p">(</span><span class="s2">&quot;Hello&quot;</span><span class="p">,</span> <span class="n">x</span><span class="p">)</span>

<span class="gp">In [43]: </span><span class="n">do</span><span class="p">(</span><span class="n">say_hello</span><span class="p">,</span> <span class="mi">4</span><span class="p">)</span>
<span class="go">Hello 4</span>
<span class="gh">Out[43]: </span><span class="go">4</span>

<span class="gp">In [44]: </span><span class="nb">list</span><span class="p">(</span><span class="nb">map</span><span class="p">(</span><span class="k">lambda</span> <span class="n">x</span><span class="p">:</span> <span class="n">do</span><span class="p">(</span><span class="n">say_hello</span><span class="p">,</span> <span class="n">x</span><span class="p">),</span> <span class="p">[</span><span class="mi">1</span><span class="p">,</span><span class="mi">2</span><span class="p">,</span><span class="mi">3</span><span class="p">]))</span>
<span class="go">                   Hello 1</span>
<span class="go">Hello 2</span>
<span class="go">Hello 3</span>
<span class="gh">Out[44]: </span><span class="go">[1, 2, 3]</span>
</pre></div>
</div>
<p>Implement a version of
this function named <code class="docutils literal"><span class="pre">my_do</span></code>.</p>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="juxt" >
    <ol start=23 class=arabic><li class="alert alert-warning">

    <p>The <code class="docutils literal"><span class="pre">toolz</span></code> library includes a function called <code class="docutils literal"><span class="pre">juxt</span></code> that takes a
sequence of functions that returns a function.  The new function will apply
each of the original functions to any input and return a tuple containing
the results.</p>
<p>For example, suppose that we want to make a function that (a)
adds 1, (b) doubles and (c) squares any input.  This is accomplished using
<code class="docutils literal"><span class="pre">juxt</span></code> as follows.</p>
<div class="highlight-ipython"><div class="highlight"><pre><span></span><span class="gp">In [45]: </span><span class="kn">from</span> <span class="nn">toolz</span> <span class="kn">import</span> <span class="n">juxt</span>

<span class="gp">In [46]: </span><span class="n">f</span> <span class="o">=</span> <span class="n">juxt</span><span class="p">(</span><span class="k">lambda</span> <span class="n">i</span><span class="p">:</span> <span class="n">i</span> <span class="o">+</span> <span class="mi">1</span><span class="p">,</span>
<span class="gp">   ....: </span>         <span class="k">lambda</span> <span class="n">i</span><span class="p">:</span> <span class="mi">2</span><span class="o">*</span><span class="n">i</span><span class="p">,</span>
<span class="gp">   ....: </span>         <span class="k">lambda</span> <span class="n">i</span><span class="p">:</span> <span class="n">i</span><span class="o">**</span><span class="mi">2</span><span class="p">)</span>
<span class="gp">   ....: </span>

<span class="gp">In [47]: </span><span class="n">f</span><span class="p">(</span><span class="mi">3</span><span class="p">)</span>
<span class="gh">Out[47]: </span><span class="go">(4, 6, 9)</span>

<span class="gp">In [48]: </span><span class="k">assert</span> <span class="n">f</span><span class="p">(</span><span class="mi">5</span><span class="p">)</span> <span class="o">==</span> <span class="p">(</span><span class="mi">5</span> <span class="o">+</span> <span class="mi">1</span><span class="p">,</span> <span class="mi">2</span><span class="o">*</span><span class="mi">5</span><span class="p">,</span> <span class="mi">5</span><span class="o">**</span><span class="mi">2</span><span class="p">)</span>
</pre></div>
</div>
<p>Implement a version of this function named <code class="docutils literal"><span class="pre">my_juxt</span></code>.</p>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="pipe" >
    <ol start=24 class=arabic><li class="alert alert-warning">

    Write and implementation of <code class="docutils literal"><span class="pre">pipe</span></code> called <code class="docutils literal"><span class="pre">my_pipe</span></code> that takes an expression
<code class="docutils literal"><span class="pre">x</span></code> as the first argument and a sequence of functions <code class="docutils literal"><span class="pre">f1</span></code>, <code class="docutils literal"><span class="pre">f2</span></code>, ... as
the remaining arguments and applies each function to the result of <code class="docutils literal"><span class="pre">x</span></code>, from
left to right.   Your function should pass the following assertions.
    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="thread_first" >
    <ol start=25 class=arabic><li class="alert alert-warning">

    Implement a version of <code class="docutils literal"><span class="pre">thread_first</span></code> called <code class="docutils literal"><span class="pre">my_thread_first</span></code> that
threads a value through the first argument of each of a list of functions.
If any of the functions take more than one argument, these arguments can be
given along with the function in a tuple.  For example, <code class="docutils literal"><span class="pre">my_thread_first(5,</span>
<span class="pre">f,</span> <span class="pre">g)</span></code> should be equivalent to <code class="docutils literal"><span class="pre">g(f(5))</span></code> and <code class="docutils literal"><span class="pre">my_thread_first(5,</span> <span class="pre">(add,</span>
<span class="pre">1),</span> <span class="pre">(pow,</span> <span class="pre">2))</span></code> expands to <code class="docutils literal"><span class="pre">pow(add(5,1),</span> <span class="pre">2)</span></code>.  Note that each expression
goes in the first argument of the next function call.
    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="files_ex_6_2" >
    <ol start=26 class=arabic><li class="alert alert-warning">

    <p>Download, unzip and extract the file titled <strong>Salaries.csv</strong> and
<strong>AllstarFull.csv</strong> from <a class="reference external" href="http://seanlahman.com/files/database/baseballdatabank-master_2016-03-02.zip">seanlahman.com</a>
.</p>
<p>Write a program that computes the average salary of all players in the all
star game in for each year between 2000 and 2010.  Write all of your
functions and expressions only using funcitonal tools like <code class="docutils literal"><span class="pre">with_iter</span></code>,
<code class="docutils literal"><span class="pre">map</span></code>, <code class="docutils literal"><span class="pre">filter</span></code>, <code class="docutils literal"><span class="pre">reduce</span></code>, <code class="docutils literal"><span class="pre">compose</span></code> and/or <code class="docutils literal"><span class="pre">pipe</span></code>.  <strong>Hint:</strong> You
will want to use <code class="docutils literal"><span class="pre">groupby</span></code> in your solution.</p>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="files_ex_8_2" >
    <ol start=27 class=arabic><li class="alert alert-warning">

    <p>Download, unzip and extract the file titled <strong>Master.csv</strong> and
<strong>BattingPost.csv</strong>.  from <a class="reference external" href="http://seanlahman.com/files/database/baseballdatabank-master_2016-03-02.zip">seanlahman.com</a></p>
<p>Write a program that computes the total runs scored in the post season by
all players from from each state (all time). Write all of your
functions and expressions only using funcitonal tools like <code class="docutils literal"><span class="pre">with_iter</span></code>,
<code class="docutils literal"><span class="pre">map</span></code>, <code class="docutils literal"><span class="pre">filter</span></code>, <code class="docutils literal"><span class="pre">reduce</span></code>, <code class="docutils literal"><span class="pre">compose</span></code> and/or <code class="docutils literal"><span class="pre">pipe</span></code>.  <strong>Hint:</strong> You
will want to use <code class="docutils literal"><span class="pre">groupby</span></code> in your solution.</p>

    </li></ol>
    </div>
    </div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - 7.8. Working with Large Files' data-toggle="tooltip">
    <a href="WorkingWithLargeFiles.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="WorkingWithLargeFiles.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Lab 1 - Closures and Scope' data-toggle="tooltip" >
    <a href="../Labs/ClosuresAndScope.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="../Labs/ClosuresAndScope.html" >Next Section - Lab 1 - Closures and Scope</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2017 Todd Iverson.
      Created using <a href="http://sphinx.pocoo.org/">Sphinx</a> 1.4.5.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>